---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```



`@instructions`
Import the scikit-learn library.

`@hint`
In python import is accomplished with > import {module_name} <

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}
import sklearn
```

`@sct`
```{python}
check_object(sklearn, missing_msg='yikes')
```
